Structure of the Calvin-Benson-Bassham sedoheptulose-1,7-bisphosphatase from the model microalga Chlamydomonas reinhardtii

Author response:

The following is the authors’ response to the previous reviews.

Recommendations for the authors:

Reviewer #1:

The authors have thoroughly changed the manuscript and addressed most of my concerns. I appreciate adding the activity assays of the C115/120S mutants, however, I suggest that the authors embed and also discuss these data more clearly. It also escaped my attention earlier that the positioning of the disulfide bond is 117-122 in the deposited PDBs instead of 115-120. The authors should carefully check which positioning is correct here.

We thank reviewer #1 for his or her careful assessment of our revised manuscript. As suggested, we detailed the results section “CrSBPase enzymatic activity” with additional numerical values, and discussed more clearly the comparisons of results for activity assays of mutants C115S and C120S in the section “Oligomeric states of CrSBPase”. Residues numbering was carefully proof-checked throughout the manuscript for correctness and homogeneity. C115 and C120 are numbered according to best databases consensus, ie. GenBank and Uniprot, and may differ from one database to another (including PDB) due to varying numbering rules. We clarified the chosen nomenclature in methods section “Cloning and mutagenesis of CrSBPase expression plasmids”.

Line 246-250: I think it is evident that the two SBPase structures superpose well given the sequence identity of more than 70%. However, it would be great to include a superposition of the two structures in Figure 1, especially with regard to the region harboring C115 and C120.

We added a panel showing superimposition of CrSBPase 7b2o and PpSBPase 5iz3 and made a close-up view around the region C115-C120 in supplementary figure 5. Given the density in information of figure 1 we prefer not to add additional images on it. Supplementary figure 5 was initially intended to illustrate sequence conservation/variation among homologs, thus fitting with the objective to compare past and present XRC results.

Line 255-266: I am again missing a panel in Figure 1 here, e.g. a side-by-side view of Xray vs AF2/3 structure.

We added another panel in supplementary figure 5 to visually compare side-by-side SBPase crystallographic structure 7b2o and our AF3 model. Again, for the sake of clarity we prefer not to overload figure 1 with additional panels. This will also enable thorough comparison of past XRC of PpSBPase, present XRC of CrSBPase, and various AF models (see below, oligomer comparisons).

Line 261-266: Did the authors predict dimers and tetramers using AF3? What are the confidence metrics in this case? Do the authors see differences to the monomer prediction in case a multimer is confidently predicted?

We modeled dimers and tetramers using AF3 and added them on supplementary figure 5 side by side with protomer of XRC model 7b2o and with monomer predicted by AF3. Color code for supplementary figure 5 panels F-H is according to AF standard representation of plDDT. Confidence metrics per residue correspond to very high reliability (navy blue) or, locally, confident prediction (cyan) and overall prediction scores range from pTM=0.85-0.91, a high-quality prediction. Interface prediction score is high for both dimer (ipTM=0.9) and tetramer (ipTM=0.82). We reported these data in supplementary figure 5 and corresponding updated legend. XRC and AF models all align with RMSD<0.5 Å, indicating a globally unchanged structure of the protomer in the various methods and oligomeric states.

Line 441: How does the oligomeric equilibrium change in C115/120S mutants? This information should be added for the mutants. Besides, the mAU units in Fig. 6 could be normalized to allow an easier comparison between the chromatograms of wt and mutants.

Change in oligomeric equilibrium is assessed by size-exclusion chromatography of WT and mutants C115S, C120S as reported in figure 6A. We made quantitative estimation of WT, and C115S and C120S mutants equilibrium by comparing maximal peak intensity and added this information in the text. Briefly, the oligomer ratio on a scale of 100 is 9:48:43 for WT, 42:25:33 for mutant C115S, and 29:17:54 for mutant C120S (ratio expressed as tetramer:dimer:monomer). We prefer not to normalize values of absorbance, but rather keep the actual measurement of absorbance at 280 nm on the chromatogram of figure 6, for the sake of consistency with the added text and for a more transparent report of the experiment.

Line 447: WT activity is 12.15+-2.15 and both mutants have a higher activity. The authors should check if their values (96% and 107%) are correct. Besides, did the authors check if the increase in C120S is statistically significant? My impression is that both mutants have a higher activity than the wildtype, in both correlating with increased fractions of the tetramer. This would also make sense, as the corresponding region is part of the tetramer interface in the crystal packing.

The reported activity values were checked for correctness. Wild-type SBPase specific activity at 12.5 ±2.15 µmol(NADPH) min-1 mg(SBPase)-1 was obtained by pre-incubating the enzyme with 1 µM CrTRXf2 supplemented with 1 mM DTT and 10 mM Mg2+, while the results of supplementary figure 14 reporting the comparison of activation of WT and mutants, with a variation of 107 or 96 %, were obtained with a slightly different protocol for pre-incubation of the enzyme with 10 mM DTT and 10 mM Mg2+. Please note that whether WT enzyme was assayed in 10 mM DTT 10 mM Mg or in 1 µM TRX 1 mM DTT 10 mM Mg, its specific activity appears equal within experimental error. Both mutants have nearly the same activity than the WT in the assay reported in supplementary figure 14: we fully agree that 107% (and 96%) variation is indeed not significant considering the uncertainty of the measurement (see error bars representing standard deviations of the mean in supplementary figure 14). We added this important information in the text. Even though both mutations stabilize the most active tetramer in untreated recombinant protein, we think that after reducting treatment both WT and mutants all reach the same maximal activity because they all form an equivalent proportion of the active tetramer versus alternative oligomeric states. We furhter interprete this piece of data as a decoupling of reduction and catalysis: in physiological conditions we assume that SBPase would initiate activation upon the reduction of disulfide bridges, including but not limited to C115-C120 that restricts the entry into fully active tetramer, at which point SBPase in reduced form reaches maximal activity until another post-translational signal eventually changes its conformation and oligomerisation.

We thank again reviewer 1 for his or her assessment and valuable suggestions.

Changes from the earlier version to the revised version

Added in the revised version: Abstract

Abstract

The Calvin-Benson-Bassham cycle (CBBC) performs carbon fixation in photosynthetic organisms. Among the eleven enzymes that participate in the pathway, sedoheptulose-1,7-bisphosphatase (SBPase) is expressed in photo-autotrophs and catalyzes the hydrolysis of sedoheptulose-1,7- bisphosphate (SBP) to sedoheptulose-7-phosphate (S7P). SBPase, along with nine other enzymes in the CBBC, contributes to the regeneration of ribulose-1,5-bisphosphate, the carbon-fixing co- substrate used by ribulose-1,5-bisphosphate carboxylase/oxygenase (Rubisco). The metabolic role of SBPase is restricted to the CBBC, and a recent study revealed that the three-dimensional structure of SBPase from the moss Physcomitrium patens was found to be similar to that of fructose-1,6- bisphosphatase (FBPase), an enzyme involved in both CBBC and neoglucogenesis. In this study we report the first structure of an SBPase from a chlorophyte, the model unicellular green microalga Chlamydomonas reinhardtii. By combining experimental and computational structural analyses, we describe the topology, conformations and quaternary structure of Chlamydomonas reinhardtii SBPase (CrSBPase). We identify active site residues and locate sites of redox- and phospho-post- translational modifications that contribute to enzymatic functions. Finally, we observe that CrSBPase adopts distinct oligomeric states that may dynamically contribute to the control of its activity.

Introduction

Photosynthetic carbon fixation is performed by eleven enzymes that collectively operate the Calvin- Benson-Bassham cycle (CBBC) (Calvin 1962, Johnson 2016), wherein Rubisco plays the prominent role in catalyzing fixation of inorganic carbon onto a pentose acceptor: ribulose-1,5-bisphosphate (RuBP). The net gain of the CBBC is one triose in the form of glyceraldehyde-3-phosphate (G3P) for every three carboxylation reactions catalyzed by Rubisco (Johnson and Alric 2013). Regeneration of RuBP is ensured by the other ten enzymes of the CBBC, among which phosphoribulokinase (PRK) and sedoheptulose-1,7-bisphosphatase (SBPase) are the only enzymes that are exclusively involved in this pathway. The remaining eight enzymes are paralogs of those involved in cytoplasmic glycolysis, gluconeogenesis or the pentose-phosphate pathway. Hence, Rubisco, PRK, and SBPase are expressed in photosynthetic organisms and fulfill biochemical functions so far uniquely attributed to the photosynthetic carbon fixation.

Pioneering studies on SBPase were carried out using the native enzyme purified from wheat chloroplasts (Woodrow 1982) and corn leaves (Nishizawa and Buchanan 1981), allowing quantitative assays of its catalytic function. This enzyme catalyzes the hydrolysis of sedoheptulose-1,7- bisphosphate (SBP) into sedoheptulose-7-phosphate (S7P) and requires magnesium ions for both catalysis and activation. In addition, SBPase activity is modulated by pH and redox switching mechanisms (Anderson 1974, Woodrow and Walker 1980, Cadet and Meunier 1988). Redox control is exerted through dithiol/disulfide interchanges dependent upon the ferredoxin-thioredoxin cascade arising from the activity of the illuminated photosystems (Schürmann and Buchanan 1975). In tobacco leaf extracts, SBPase reaches its maximal activity after 10 minutes illumination (Zimmer, Swart et al. 2021), and this activation kinetics is consistent with the slow reduction pattern of Arabidopsis SBPase in vitro (Yoshida and Hisabori 2018).

The crystal structure of SBPase from the moss Physcomitrium patens (PpSBPase) has been solved at a resolution of 1.3 Å (Protein Data Bank entry : 5iz3) (Gutle, Roret et al. 2016). PpSBPase fold is highly similar to that of moss fructose-1,6-bisphosphatase (FBPase), yet the protein was shown to distinctively oligomerize as a homodimer and to form a disulfide bridge in a different motif than that found in plant FBPase (Chiadmi, Navaza et al. 1999) (Gütle, Roret et al. 2016). In spite of their homology, the sequences of FBPase and SBPase seem to have evolved different regulatory and catalytic properties that allow different contributions to the CBBC. Indeed, both enzymes specifically recognize their respective 6-carbon and 7-carbon sugar substrates (i.e., fructose-1,6-bisphosphate and SBP, respectively), with only SBPase capable to hydrolyze the other substrate with a resulting catalysis lower than that measured with SBP (Gütle, Roret et al. 2016). To date, the structural components contributing to active site selectivity (i.e., substrate recognition) can be retrieved from two studies reporting the crystal structure of the dual function fructose-1,6-bisphosphatase/sedoheptulose-1,7-bisphosphatase (F/SBPase) from Thermosynechococcus elongatus in complex with SBP (Cotton, Kabasakal et al. 2015) and from Synechocystis sp. PCC 6803 in complex with FBP (PDB entry 3RPL). Recently, it was proposed that the CBBC can bypass fructose-1,6- bisphosphate (FBP) altogether, by only relying on SBP, glycolysis and the oxidative pentose- phosphate pathway paralogs to regenerate the substrate of Rubisco (Ohta 2022).

Several studies have reported that SBPase, along with Rubisco, FBPase, aldolase and transketolase, constitute a metabolic bottleneck for the CBBC ((Raines 2022) and references therein). Overexpression of SBPase was found to effectively enhance photosynthesis and growth in plants (Lefebvre, Lawson et al. 2005, Feng, Han et al. 2007, Feng, Wang et al. 2007, Rosenthal, Locke et al. 2011, Simkin, Lopez-Calcagno et al. 2017) and in the green microalga Chlamydomonas reinhardtii (Hammel, Sommer et al. 2020), indicating that SBPase can be considered as a limiting factor for the photosynthetic carbon fixation. Based on these evidences, it is plausible to consider that SBPase is a good target to improve the photosynthetic capacity by means of metabolic engineering coupled to synthetic biology (Kubis and Bar-Even 2019), especially in the experimentally convenient multi-omic model alga Chlamydomonas (Mettler, Muhlhaus et al. 2014, Schmollinger, Mühlhaus et al. 2014).

In Chlamydomonas, SBPase is encoded by a single nuclear gene (Cre03.g185550), and the protein (Uniprot entry P46284) is addressed to the chloroplast by a 54-residue amino-terminal peptide (Emanuelsson, Brunak et al. 2007, Almagro Armenteros, Salvatore et al. 2019). SBPase abundance in the algal cell was quantified by mass spectrometry as 1.1-1.3 ± 0.3 amol/cell, representing 0.15% of total cell proteins (Hammel, Sommer et al. 2020). In Chlamydomonas protein extracts, SBPase was shown to undergo post-translational modifications by phosphorylation of serines and threonines (Werth, McConnell et al. 2017), and redox modifications of cysteine thiols (Zaffagnini, Bedhomme et al. 2012, Morisse, Zaffagnini et al. 2014, Pérez-Pérez, Mauriès et al. 2017). However, the mechanisms by which post-translational modifications affect the catalytic activity of CrSBPase is still lacking. Here, we describe the crystal structures of CrSBPase in two redox states, examine the catalytic activity of recombinant CrSBPase, and investigate the correlation between the redox state and the oligomerization of the protein. Our results suggest that a redox control over oligomer equilibrium could represent a novel mechanism to regulate enzymatic activities in the CBBC.

Material and methods

Chemical reagents were purchased from Merck (Darmstadt, Germany), chromatography material was provided by Cytiva (Vélizy-Villacoublay, France), crystallization solutions and consumable were purchased from Qiagen (Hilden, Germany), Hampton Research (Aliso Viejo, CA USA), and SPT Labtech (Melbourn, UK).

Cloning and mutagenesis of CrSBPase expression plasmids

Earlier version: The sequence coding for full-length mature SBPase was amplified by reverse transcriptase polymerase chain reaction (RT-PCR) from Chlamydomonas reinhardtii total ribonucleic acids extract and inserted at sites 5’-NdeI and BamHI-3’ into expression vector pET3c modified for the fusion of an amino-terminal hexa-histidine tag in the recombinant protein, yielding plasmid pSL-175-His6-CrSBPase-WT. PCR primers were 5’-TCTCGCCATATGGCCGTTCTGACCCAGGCC-3’ (forward, NdeI site underlined) and 5’-CGGGTGGGATCCTTAGGCAGCCACCTTCTC-3’ (reverse, BamHI site underlined). The encoded protein encompasses 331 amino acids, with a molar mass of 36,176.7 Da and a theoretical extinction coefficient of 24,722.5 M-1.cm-1. Unless otherwise stated, amino acid residues numbering is that of GenBank entry DAA79954.1. Point mutants of cysteines 115 and 120 singly substituted with serines were generated by site-directed PCR mutagenesis with primers listed in supplementary table 1, producing plasmids pSL-181-His6-CrSBPase-C115S, and pSL-182-His6-CrSBPase-C120S.

Revised version: The sequence coding for full-length mature SBPase was amplified by reverse transcriptase polymerase chain reaction (RT-PCR) from Chlamydomonas reinhardtii total ribonucleic acids extract and inserted at sites 5’-NdeI and BamHI-3’ into expression vector pET3c modified for the fusion of an amino-terminal hexa-histidine tag in the recombinant protein, yielding plasmid pSL-175-His6- CrSBPase-WT. PCR primers were 5’-TCTCGCCATATGGCCGTTCTGACCCAGGCC-3’ (forward, NdeI site underlined) and 5’-CGGGTGGGATCCTTAGGCAGCCACCTTCTC-3’ (reverse, BamHI site underlined). The encoded protein encompasses 331 amino acids, with a molar mass of 36,176.7 Da and a theoretical extinction coefficient of 24,722.5 M-1.cm-1. Unless otherwise stated, amino acid residues numbering throughout the manuscript is that of GenBank entry DAA79954.1 and Uniprot entries A8IRK4 and P46284. Point mutants of cysteines 115 and 120 singly substituted with serines were generated by site-directed PCR mutagenesis with primers listed in supplementary table 1, producing plasmids pSL- 181-His6-CrSBPase-C115S, and pSL-182-His6-CrSBPase-C120S.

Recombinant CrSBPase purification

Plasmids pSL-175, -181, and -182 were used to transform Escherichia coli BL21(DE3). The transformants were grown at 37 °C in 2YT supplemented with 150 µg/mL ampicillin. When cultures reached an OD600 of ∼0.6, T7-dependent CrSBPase overexpression was induced by the addition of isopropyl-β-D-thiogalactopyranoside (IPTG) at 0.2 mM for 3 h.

Harvested cell pellets were resuspended in 20 mM Tris pH 7.9, 100 mM NaCl (buffer A) and lyzed by 1 sec / 1 sec pulsed sonication for 3 min. The soluble fraction of the lysate was separated by 20 min centrifugation at 20,000 g and loaded on 2 mL Ni-NTA resin for His-tag mediated affinity chromatography. The resin was washed in four steps with 25 mL buffer A, 25 mL buffer A supplemented with 10, 20, and 30 mM imidazole, and CrSBPase was eluted in three steps with 10 mL buffer A supplemented with 100, 200, and 300 mM imidazole. Pooled eluates containing CrSBPase were loaded on HiLoad Superdex 200 26/600 size-exclusion chromatography column and eluted in buffer A. Protein purity was assessed by electrophoresis on 12% acrylamide gel under denaturing and reducing conditions (supplementary figure 1). Peak fractions were assembled and concentrated by ultrafiltration on Amicon units of MWCO 30,000 Da cut-off. Final concentrations of purified proteins were in the 1-10 mg/ml range as measured by Nanodrop ND-2000 spectrophotometer using an ε280 of 24,722.5 M-1.cm-1.

Western blot

Polyclonal antibodies were generated by immunization of rabbits with pure recombinant CrSBPase (Genecust, Boynes France). Analyzed fractions were separated by denaturing polyacrylamide electrophoresis and transferred to 0.2 µm nitrocellulose membrane for detection with primary antibody subsequently detected by secondary anti-rabbit antibodies coupled to horseradish peroxidase. Detection was done with commercial ECL peroxidase assay (GE Healthcare, Chicago IL USA) with a Chemidoc (Bio-Rad, Hercules CA USA).

Crystallization and structure determination

Earlier version: Pure recombinant His6-CrSBPase-WT concentrated at 3-6 mg/mL was tested for crystallization on screens JCSG I-IV in 200 nL sitting drops. Single crystals were obtained and optimized with mother liquor 0.1 M sodium HEPES pH=7.5, 2% (v/v) polyethylene glycol 400, 2.0 M ammonium sulfate, flash-frozen in mother liquor supplemented with 25% glycerol and tested for X-ray diffraction. Complete X-ray diffraction datasets were collected at SOLEIL beamline Proxima-2A and allowed the determination of the crystal structure by model replacement with 5iz3, the structure of ortholog SBPase from the moss Physcomitrium patens. Model building in COOT and refinement in the PHENIX software suite yielded structure deposited under PDB ID 7b2o (Emsley, Lohkamp et al. 2010) (Adams, Afonine et al. 2011) (Liebschner, Afonine et al. 2019). After we obtained our experimental structure, the European Bioinformatics Institute (Hinxton, UK) communicated high-accuracy prediction models computed by Deepmind ALPHAFOLD2 (Jumper, Evans et al. 2021), including a prediction of CrSBPase structure (https://alphafold.ebi.ac.uk/entry/P46284) that matches chain A in 7b2o crystal structure with a root mean square deviation (RMSD) of 0.453 Å. Hence, X-ray crystallography and AlphaFold independently contribute to the proposed structural analysis.

Revised version: Pure recombinant His6-CrSBPase-WT concentrated at 3-6 mg/mL was tested for crystallization on screens JCSG I-IV in 200 nL sitting drops. Single crystals were obtained and optimized with mother liquor 0.1 M sodium HEPES pH=7.5, 2% (v/v) polyethylene glycol 400, 2.0 M ammonium sulfate, flash- frozen in mother liquor supplemented with 25% glycerol and tested for X-ray diffraction. Complete X- ray diffraction datasets were collected at SOLEIL beamline Proxima-2A and allowed the determination of the crystal structure by model replacement with 5iz3, the structure of ortholog SBPase from the moss Physcomitrium patens.Model building in COOT and refinement in the PHENIX software suite yielded structure deposited under PDB ID 7b2o (Emsley, Lohkamp et al. 2010) (Adams, Afonine et al. 2011) (Liebschner, Afonine et al. 2019). After we obtained our experimental structure, the European Bioinformatics Institute (Hinxton, UK) communicated high-accuracy prediction models computed by Deepmind AlphaFold 2 (Jumper, Evans et al. 2021), including a prediction of CrSBPase structure (https://alphafold.ebi.ac.uk/entry/P46284) that matches chain A in 7b2o crystal structure with a root mean square deviation (RMSD) of 0.453 Å. Hence, X-ray crystallography and AlphaFold independently contribute to the proposed structural analysis.

In order to obtain the structure of CrSBPase in a reduced state, crystals were grown in the presence of 10 mM of the reducing agent tris-(2-carboxyethyl)phosphine (TCEP) in condition JCSG IV E8/56 0.2 M lithium sulfate, 0.1 M Tris pH 7.0, and 1.0 M sodium/potassium tartrate. Crystals were cryo- protected in mother liquor supplemented with 25% ethylene glycol, flash-frozen and tested for X-ray diffraction at SOLEIL beamline Proxima-2A. A complete dataset was collected that allowed the determination of the structure by molecular replacement with 7b2o as a template, resulting in structure 7zuv. Among the eight CrSBPase copies of the asymmetric unit, two had their disulfides absent as a consequence of the TCEP treatment while the others six subunits were essentially identical to untreated structures in entry 7b2o (i.e., fully oxidized form).

Enzymatic assays

The catalytic activity of 10-200 nM pure recombinant CrSBPase was spectrophotometrically assayed in buffer A by coupling the hydrolysis of 1 mM FBP into F6P to the reduction of 0.5 mM NADP+ into NADPH through isomerization of F6P into glucose-6-phosphate (G6P) by 0.5 U.mL-1 phosphoglucose isomerase (PGI) and the oxidation of G6P into 6-phosphogluconolactone (6PGL) by 0.1-0.5 U.mL-1 glucose-6-phosphate dehydrogenase (G6PDH). Reporter assay is coupled in a molar ratio FBP:F6P:G6P:6PGL of 1:1:1:1 to the oxidation of one molar equivalent of NADP+ into NADPH that we recorded over 20 minutes by measuring the absorbance at 340 nm using a UVIKON spectrophotometer. Pre-treatments with 0-100 mM reduced dithiothreitol (DTTred), 0-30 mM MgSO4, and 0-15 µM recombinant CrTRX-f2 supplemented with 1 mM DTTred (Lemaire, Tedesco et al. 2018) were carried out for 30 min at room temperature before measuring enzyme activity as described above.

Titration of accessible cysteine thiols

Earlier version: Quantification of free reactive cysteines on recombinant CrSBPase was performed with 5,5ʹ-dithio-bis-(2-nitrobenzoic acid) (DTNB, Ellman’s reagent) mixed with the purified recombinant protein in a 200:1 molar ratio in a 1 mL reaction volume containing 100 mM Tris-HCl, pH 7.9 as described previously (Meloni, Fanti et al. 2024). Reaction of DTNB with solvent-exposed cysteine thiols produces 2-nitro-5-thiobenzoate (TNB-) that we quantified by monitoring the absorbance at 412 nm and a molar extinction coefficient of TNB- of 14.15 mM-1.cm-1. We calculated the number of free reactive cysteines per CrSBPase monomer by dividing the molar concentration of TNB- by that of the protein.

Revised version: Quantification of free reactive cysteines on recombinant CrSBPase was performed with 5,51-dithio- bis-(2-nitrobenzoic acid) (DTNB, Ellman’s reagent) mixed with the purified recombinant protein in a 200:1 molar ratio in a 1 mL reaction volume containing 100 mM Tris-HCl, pH 7.9 as described previously (Meloni, Fanti et al. 2024). Reaction of DTNB with solvent-exposed cysteine thiols produces 2-nitro-5-thiobenzoate (TNB-) that we quantified by monitoring the absorbance at 412 nm and a molar extinction coefficient of TNB- of 14.15 mM-1.cm-1. We calculated the number of free reactive cysteines per CrSBPase monomer by dividing the molar concentration of TNB- by that of the protein.

Molecular dynamics

Molecular dynamics (MD) simulations were performed with the Gromacs package (v2021.4). A dimer of CrSBPase was solvated in a rhombic dodecahedron box with at least 8 Å between the protein atoms and the box edges, which corresponds to ∼64000 atoms in total (∼18000 water molecules) and a triclinic box of 97x97x68 Å3. The box was neutralized with 18 sodium atoms to make it neutral. The starting conformation was the crystallized structure of oxidized CrSBPase, and we protonated C115 and C120 to manually create the reduced state. The protein was described with the Amber14SB force field and water with the TIP3P model. Non-bonding interactions were described with standard parameters: van der Waals interactions had a cut-off of 8 Å, and electrostatic interactions were computed with PME with default parameters with a separation between spaces at 8 Å. Bonds containing a hydrogen atom were restrained with the LINCS algorithm with default parameters. After an energy minimization, the system was equilibrated under the NPT conditions with a δt=1 fs timestep, the velocity-rescale thermostat and the Berendsen barostat. Velocities were generated at 100 K and the system was heated up to 300 K in 400 ps before performing 100 ps at 300 K. We then ran simulations of 20 µs under the NPT conditions with a δt=2 fs timestep, the velocity-rescale thermostat and the Parrinello-Rahman barostat.

Cultivation of Chlamydomonas for native SBPase analysis

Chlamydomonas reinhardtii strain D66 was grown in 50 mL Tris-acetate-phosphate (TAP) medium under 40 μE.m-2.s-1 continuous light until density reached 5x104 cell.mL-1, harvested by 10 min centrifugation at 3,000 g. Cell pellet was resuspended in 100 mM Tris-HCl, pH 7.9, 100 mM NaCl, 5 mM L-cystéine, 500 μM NADH, and 1 mM EDTA (buffer L) and lysed by 3 freeze-thaw cycles in liquid nitrogen and water bath. Soluble fraction was separated by 10 min centrifugation at 30,000 g and loaded for size-exclusion chromatography on a calibrated Superose 6 16/600 column equilibrated in buffer L. Aliquots of 500 µL elution fractions were analyzed by western blot with primary antibodies raised in rabbits against pure recombinant CrSBPase (Covalab, Bron, France) and secondary anti- rabbit IgG coupled to peroxidase (Sigma-Aldrich reference A9169, Saint Louis, USA) revealed by ECL Prime (GE Healthcare).

Small-angle X-ray scattering

Pure recombinant CrSBPase was injected on Agilent HPLC BioSEC-3 300 column in buffer A, in line with the X-ray beam capillary at SOLEIL beamline SWING. CrSBPase was untreated or pre-treated with 10 mM of reducing agent dithiothreitol (DTTred). Diffusion curves were analyzed with ATSAS software to determine the molecular mass and oligomeric state of the protein.

Results and Discussion

Crystal structure of CrSBPase

Earlier version: Pure recombinant CrSBPase crystallized in space group P21212 and provided a complete X-ray diffraction dataset at a resolution of 3.09 Å. The asymmetric unit is composed of six polypeptide chains packing as three dimers and four water molecules could be modeled. All chains align within a root mean square deviation (RMSD) <0.25 Å and can be considered identical. 309 out of the 331 residues were built in the electron density, the unmodelled sequences are in the disordered extensions at the amino- and carboxy-termini of the model. CrSBPase folds into two distinct domains (figure 1A-B). The amino-terminal domain (NTD, residues 73-247) is composed of helices 1, 2, and 3 and a mixed β-sheet formed by strands 3-9. The carboxy-terminal domain (CTD, residues 248-383) is composed of helices 4, 5, and 8 and helices 6 and 7 that sandwich a mixed β-sheet formed by strands 10-13. α-helices 1-3 pack on the amino-terminal end of the seven-strand sheet, while α-helices 4,5, and 8 pack on the carboxy-terminal end of the four-strand sheet. α-helices 6 and 7 are sandwiched in between the two sheets. CrSBPase chain A aligns with previously reported crystal structure of PpSBPase with an RMSD of 0.478 Å (Gütle et al. 2018 PDB 5IZ3). This close structural similarity between orthologues indicates common structural properties supporting a conservation of the SBPase function in the evolutionary time separating the microalga Chlamydomonas from the land moss P. patens. A short antiparallel β-sheet is constituted of 4-residue strands 1 and 2 separated by a β-turn and preceded by loop A113SCAGTAC120. This motif projects to solvent from α-helices 2 and 3 (figure 1C-D). Residues 113-130 forming the loop and β-hairpin (LBH) motif are positioned away from the core of the protein, at 15.3 Å from water molecule 401 (W401), and will be further discussed in the next sections.

Revised version: Pure recombinant CrSBPase crystallized in space group P21212 and provided a complete X-ray diffraction dataset at a resolution of 3.09 Å. The asymmetric unit is composed of six polypeptide chains packing as three dimers and four water molecules could be modeled. All chains align within a root mean square deviation (RMSD) <0.25 Å and can be considered identical. 309 out of the 331 residues were built in the electron density, the unmodelled sequences are in the disordered extensions at the amino- and carboxy-termini of the model. CrSBPase folds into two distinct domains (figure 1A-B). The amino-terminal domain (NTD, residues 73-247) is composed of helices 1, 2, and 3 and a mixed β-sheet formed by strands 3-9. The carboxy-terminal domain (CTD, residues 248-383) is composed of helices 4, 5, and 8 and helices 6 and 7 that sandwich a mixed β-sheet formed by strands 10-13. α-helices 1-3 pack on the amino-terminal end of the seven-strand sheet, while α-helices 4,5, and 8 pack on the carboxy-terminal end of the four-strand sheet. α-helices 6 and 7 are sandwiched in between the two sheets. CrSBPase chain A aligns with previously reported crystal structure of PpSBPase with an RMSD of 0.478 Å (Gütle et al. 2018 PDB 5iz3) (supplementary figure 5C). This close structural similarity between orthologues indicates common structural properties supporting a conservation of the SBPase function in the evolutionary time separating the microalga Chlamydomonas from the land moss P. patens . A short antiparallel β-sheet is constituted of 4- residue strands 1 and 2 separated by a β-turn and preceded by loop A113SCAGTAC120. This motif projects to solvent from α-helices 2 and 3 (figure 1C-D). Residues 113-130 forming the loop and β- hairpin (LBH) motif are positioned away from the core of the protein, at 15.3 Å from water molecule 401 (W401), and will be further discussed in the next sections.

Added in the revised version: Crystal structure of CrSBPase.

Earlier version: Crystal structure of CrSBPase. (A) Topology is displayed as cartoon with main chain colored from blue (amino-terminus) to red (carboxy-terminus). (B) Rotated view of A. by x-axis over -90°. (C) Close-up view of region A113SCAGTAC120 with disulfide bond shown in sticks. (D) Rotated view of C. by x-axis over -90°. (E) Putative active site residues inferred from alignment with FBPase bound to FBP (5l0a, human muscle fructose-1,6-bisphosphatase E69Q mutant in active R-state in complex with fructose-1,6-bisphosphate) are represented in sticks. Residue numbering is according to Chlamydomonas SBPase Uniprot entry P46284. (F) Surface representation in the B. orientation with putative active site residues colored in cyan. Water molecule 401 (W401) oxygen is represented as a red sphere. (G) Cysteines site chains are represented in spheres. (H) Threonines and serines side chains reported to be the target of phosphorylations are represented in spheres.

Revised version: (A) Topology is displayed as cartoon with main chain colored from blue (amino-terminus) to red (carboxy-terminus). (B) Rotated view of A. by x-axis over -90°. (C) Close-up view of region A113SCAGTAC120 with disulfide bond shown in sticks. (D) Rotated view of C. by x-axis over -90°. (E) Putative active site residues inferred from alignment with FBPase bound to FBP (5l0a, human muscle fructose-1,6-bisphosphatase E69Q mutant in active R-state in complex with fructose-1,6- bisphosphate) are represented in sticks. Residue numbering is according to Chlamydomonas SBPase Uniprot entry P46284. (F) Surface representation in the B. orientation with putative active site residues colored in cyan. Water molecule 401 (W401) oxygen is represented as a red sphere. (G) Cysteines site chains are represented in spheres. (H) Threonines and serines side chains reported to be the target of phosphorylations are represented in spheres.

AF2 entry AF-P46284-F1-model_v4 from AlphaFold Protein Structure Database aligns with our crystal structure 7b2o chain E with RMSD = 0.434 Å, showing excellent agreement between experiment and prediction at the level of protein main chain. Exceptions are in local differences in several loops conformations and in the N-ter and C-ter borders of secondary structure elements. Numerous amino acid residues side chains adopt moderately distinct orientations between the computational model and the experimental structure.

Earlier version: AF3 was recently communicated (Abramson, Adler et al. 2024) along with its online prediction server hosted at https://golgi.sandbox.google.com/. CrSBPase model from AF3 align to our crystal structure 7B2O chain A with RMSD = 0.489 Å showing again their strong similarity and with a small discrepancy between AF2 and AF3 models with RMSD = 0.216 Å. The only significant deviations between 7B2O and AF3 are in the orientation of several side chains and notably on the conformation of region 114-131 that contain the LBH motif.

Revised version: AF3 was recently communicated (Abramson, Adler et al. 2024) along with its online prediction server hosted at https://golgi.sandbox.google.com/. CrSBPase model from AF3 align to our crystal structure 7B2O chain A with RMSD = 0.441 Å showing again their strong similarity (supplementary figure 5E,F) and with a small discrepancy between AF2 and AF3 models with RMSD = 0.247 Å. The only significant deviations between 7B2O and AF3 are in the orientation of several side chains and notably on the conformation of region 114-131 that contain the LBH motif.

CrSBPase catalytic pocket

Earlier version: CrSBPase is structurally similar to its ortholog from the moss Physcomitrium patens (RMSD=0.478 Å) (Gutle, Roret et al. 2016). SBPase is a homolog of FBPase, another enzyme involved in the CBBC in photosynthetic organisms and also ubiquitously involved in neoglucogenesis. Cytosolic isoforms of FBPase from non-plant sources have been extensively studied (Moorhead, Hodgson et al. 1994, Lee and Hahn 2003). For instance, the crystal structure of human muscle FBPase complexed with its substrate (FBP) was solved (PDB entry: 5l0a, to be published). We aligned the structure CrSBPase onto this human FBPase:FBP complex from entry 5l0a, confirming their similarity (RMSD = 0.824 Å) with the notable exception of the LBH motif which is absent in FBPase. In a 4-Å distance to FBPase co-crystallized with FBP we aligned 12 residues that are likely available for FBP or SBP binding by CrSBPase (figure 1E): E155, D173, D176, G177, Y287, G289, G290, M291, K317, L318, R319, and E323. These residues are exposed to the solvent in a continuous surface (figure 1F). W401 could be modelled at the center of the putative SBPase catalytic pocket, in the vicinity of E155, D173, D176, R319, and E323, in the catalytic pocket predicted from alignment with human FBPase bound to FBP. W401 specifically is in bonding distance to catalytic residues D178, D181, and E328. W401 is likely to hydrolyze the substrate SBP.

Revised version: CrSBPase is structurally similar to its ortholog from the moss Physcomitrium patens (RMSD=0.478 Å) (Gutle, Roret et al. 2016). SBPase is a homolog of FBPase, another enzyme involved in the CBBC in photosynthetic organisms and also ubiquitously involved in neoglucogenesis. Cytosolic isoforms of FBPase from non-plant sources have been extensively studied (Moorhead, Hodgson et al. 1994, Lee and Hahn 2003). For instance, the crystal structure of human muscle FBPase complexed with its substrate (FBP) was solved (PDB entry: 5l0a, to be published). We aligned the structure CrSBPase onto this human FBPase:FBP complex from entry 5l0a, confirming their similarity (RMSD = 0.824 Å) with the notable exception of the LBH motif which is absent in FBPase. In a 4-Å distance to FBPase co-crystallized with FBP we aligned 12 residues that are likely available for FBP or SBP binding by CrSBPase (figure 1E): E155, D173, D176, G177, Y287, G289, G290, M291, K317, L318, R319, and E323.

Added in the revised version: These residues are exposed to the solvent in a continuous surface (figure 1F). W401 could be modelled at the center of the putative SBPase catalytic pocket, in the vicinity of E155, D173, D176, R319, and E323, in the catalytic pocket predicted from alignment with human FBPase bound to FBP. W401 specifically is in bonding distance to catalytic residues D178, D181, and E328. W401 is likely to hydrolyze the substrate SBP.

Sites of redox post-translational modifications

Ten cysteines could be located in each monomer of CrSBPase (C115, C120, C149, C153, C169, C222, C231, C340, C355, C362). Previous redox-based proteomic studies exploiting Chlamydomonas cells or protein extract established that CrSBPase undergoes redox modifications by thioredoxin-dependent dithiol/disulfide exchange (Lemaire, Guillon et al. 2004, Pérez-Pérez, Mauriès et al. 2017),S- glutathionylation (Zaffagnini, Bedhomme et al. 2012), or S-nitrosylation (Morisse, Zaffagnini et al. 2014). While the putative Cys sites for S-glutathionylation are not known, thioredoxin and S- nitrosylation Cys targets were identified as being C222/C231/C355/C362, and C355/C362, respectively. DTNB-based thiol titration on the native purified SBPase indicates that 4.0 ± 0.6 cysteines are likely reactive owing to their relative exposure to solvent in a free thiol form. Amino acid residue accessible surface area and accessibility (ASA) calculation is consistent with DTNB titration, with C149, C153, C231, and C355 exposed to the solvent (ASA respective areas: 0.154 Å², 0.113 Å², 0.308 Å², and 10.083 Å²), and the six others either shielded from solvent (C169, C222, C340, and C362) or engaged in a disulfide bridge (C115, C120) (Ahmad, Gromiha et al. 2004). In CrSBPase crystal structure, C222 and C231 are positioned on neighboring strands of the amino-terminal domain, exposing their thiol 3.5 Å away from each other (figure 1G, supplementary figures 2A-B). The formation of an intramolecular disulfide bridge at this cysteines pair is possible if local main chain rearrangements allow a closer contact. C362 thiol is positioned 10.5 Å away from the nearest thiol group of C340. Additionally, C340 and C362 are exposed on an opposite side of the carboxy-terminal β-sheet, disfavoring the formation of an internal disulfide bridge. C355 points its side chain towards C153 and C169 with a thiol-thiol distance of 4.6 Å and 10.1 Å, respectively. Although not observed in our model, a C355-C153/C169 dithiol-disulfide exchange is thus plausible, as described for plant FBPase (Jacquot, Lopez-Jaramillo et al. 1997). Local rearrangements that occur upon FBPase disulfide reduction causes a larger scale relaxation of the active site required to reach maximal catalytic potential (Chiadmi, Navaza et al. 1999) as kinetically characterized for enzymes from Arabidopsis thaliana (Yoshida and Hisabori 2018) and Nicotiana benthamiana (Zimmer, Swart et al. 2021). In Chlamydomonas and Physcomitrium SBPases, the FBPase-type regulatory disulfide is not observed but the SBPase-type C115-C120 bridge is conserved in a nearly identical position in the LBH motif where it may regulate the enzyme by alternative conformational rearrangements. Indeed, the disulfide-bonded cysteines constrain Chlamydomonas SBPase loop A113SCAGTAC120 into a compact lasso that could open into a relaxed conformation hypothetically activating catalysis.

Sites of post-translational phosphorylations

CrSBPase was reported to be phosphorylated from algae extracts at residues T112, S114, T118, S123, T310, S311, and T313 (Wagner, Gessner et al. 2006, Werth, McConnell et al. 2017, McConnell, Werth et al. 2018). Consistently with the proximity of these residues in two sectors of the sequence (supplementary figure 3), the candidate modification sites are located in two contiguous regions in the three-dimensional structure (figure 1H). T112, S114, T118, and S123 are exposed to the solvent of the amino-terminal domain. The residues belong to or are located nearby the A113SCAGTAC120 motif where phosphorylations may cross-signal with redox modifications on C115 and C120. T310, S311, and T313 are located on the other edge of the enzyme, at the solvent exposed tip of the carboxy-terminal domain. Their co-localization at the T310SPT313 motif likely facilitates a coordinated phosphorylation of the three residues by the same kinase. Because both groups of potential phospho-sites are located at 17-18 Å from the catalytic pocket (W401 taken as a reference point), there is no straightforward mechanism by which such modifications would exert a control over CrSBPase activity. Analysis of phospho-mimicking mutants along with the identification of the actual kinases/phosphatases couples that control CrSBPase phosphorylation state would allow to comprehend the role of this putative regulatory mechanism.

CrSBPase enzymatic activity

Earlier version: Recombinant CrSBPase was assayed for its capacity to catalyze FBP hydrolysis into F6P (Gutle, Roret et al. 2016), a proxy for the hydrolysis of SBP into S7P that we could not test by lack of available substrate and kinetic reporter method. Previous studies demonstrated that plant SBPase requires the cofactor Mg2+ and a chemical reductant (DTTred) to be fully activated in vitro, reaching a KM value of 0.05 mM and a kcat of 81 sec-1 in optimal conditions (Anderson 1974, Woodrow and Walker 1980, Cadet and Meunier 1988). In agreement with this, we found a total absence of activity in the untreated CrSBPase (data not shown), whereas pre-incubation of the enzyme with DTTred and Mg2+ resulted in a strong activation (figure 2). By testing different concentrations of MgSO4 (3, 10, 15, and 20 mM) and DTTred (1, 10, 25, 50, and 100 mM), we established that 10 mM of DTT and 10 mM MgSO4 are the optimal conditions to attain the maximum specific activity (figure 2A and 2B). Unexpectedly, DTTred concentrations above 10 mM had an inhibitory effect, causing a drastic drop in activity that probably resulted from structural effects leading to denaturation of the enzyme (figure 2A). In contrast, we observed no such inhibitory effect in the presence of magnesium at concentrations above 10 mM (figure 2B). Under physiological conditions, the light-dependent redox regulation of plant SBPase is specifically exerted by f-class thioredoxins (Gutle, Roret et al. 2016), which are reduced by the ferredoxin-thioredoxin cascade (Schürmann and Buchanan 1975). The chloroplast TRX-f2 from Chlamydomonas reinhardtii was thus tested and validated for its capacity to activate CrSBPase by reduction (figure 2C). The maximum activity, corresponding to 12.15 ± 2.15 µmol(NADPH) min-1 mg(SBPase)-1, was obtained by pre-incubating the enzyme with 1 µM CrTRX-f2 supplemented with 1 mM DTT and 10 mM MgSO4, and maintained even at higher CrTRX-f2 concentrations (i.e., 5 and 10 µM) (figure 2C). CrSBPase specific activity that we report is comparable to that reported for spinach SBPase with FBP as substrate (3.5 µmol min-1 mg-1 (Seuter, Busch et al. 2002)) and to that reported for another recombinant Chlamydomonas SBPase with SBP as substrate (7.15 ± 0.61 µmol min-1 mg-1 (Tamoi, Nagaoka et al. 2005)).

Revised version: Recombinant CrSBPase was assayed for its capacity to catalyze FBP hydrolysis into F6P (Gutle, Roret et al. 2016), a proxy for the hydrolysis of SBP into S7P that we could not test by lack of available substrate and kinetic reporter method. Previous studies demonstrated that plant SBPase requires the cofactor Mg2+ and a chemical reductant (DTTred) to be fully activated in vitro, reaching a K value of 0.05 mM and a k of 81 sec-1 in optimal conditions (Anderson 1974, Woodrow and Walker 1980, Cadet and Meunier 1988). In agreement with this, we found a total absence of activity in the untreated CrSBPase (data not shown), whereas pre-incubation of the enzyme with DTTred and Mg2+ resulted in a strong activation, eg. a detectable activity at DTT ≥1 mM (figure 2A) and a stark increase in activity at MgSO4 ≥10 mM (figure 2B). By testing different concentrations of MgSO4 (3, 10, 15, and 20 mM) and DTTred (1, 10, 25, 50, and 100 mM), we established that 10 mM of DTT and 10 mM MgSO4 are the optimal conditions to attain the maximum specific activity (figure 2A and 2B). Unexpectedly, DTTred concentrations above 10 mM had a dose-dependent inhibitory effect, causing a drastic drop in activity that probably resulted from structural effects leading to denaturation of the enzyme (figure 2A, DTT ≥25 mM). In contrast, we observed no such inhibitory effect in the presence of magnesium at concentrations above 10 mM (figure 2B). Under physiological conditions, the light- dependent redox regulation of plant SBPase is specifically exerted by f-type thioredoxins (Gutle, Roret et al. 2016), which are reduced by the ferredoxin-thioredoxin cascade (Schürmann and Buchanan 1975). The chloroplast TRX-f2 from Chlamydomonas reinhardtii was thus tested and validated for its capacity to activate CrSBPase by reduction, with maximal activation reached at TRX- f2 ≥1 µM and no apparent loss of activity at higher TRX-f2 concentrations (figure 2C). The maximum SBPase reported activity, corresponding to 12.15 ± 2.15 µmol(NADPH) min-1 mg(SBPase)-1, was obtained by pre-incubating the enzyme with 1 µM CrTRX-f2 supplemented with 1 mM DTT and 10 mM MgSO4, and maintained even at higher CrTRX-f2 concentrations (i.e., 5 and 10 µM) (figure 2C). CrSBPase specific activity that we report is comparable to that reported for spinach SBPase with FBP as substrate (3.5 µmol min-1 mg-1 (Seuter, Busch et al. 2002)) and to that reported for another recombinant Chlamydomonas SBPase with SBP as substrate (7.15 ± 0.61 µmol min-1 mg-1 (Tamoi, Nagaoka et al. 2005)). We hence conclude that our assay faithfully recapitulate the level of CrSBPase hydrolase function in its reduced state.

Added in the revised version: Functional characterization of CrSBPase in vitro.

Earlier version: Functional characterization of CrSBPase in vitro. Reported enzymatic activity of CrSBPase assayed for (A) reduction by reduced dithiothreitol (DTTred) and 10 mM of MgSO4, (B) Magnesium sulfate (MgSO4) and 10 mM DTTred, and (C) recombinant thioredoxin f2 from Chlamydomonas reinhardtii (CrTRX-f2) and 10 mM MgSO4, 1 mM DTT.

Revised version: Reported enzymatic activity of CrSBPase assayed for (A) reduction by reduced dithiothreitol (DTTred) and 10 mM of MgSO4, (B) Magnesium sulfate (MgSO4) and 10 mM DTTred, and (C) recombinant thioredoxin f2 from Chlamydomonas reinhardtii (CrTRX-f2) and 10 mM MgSO4, 1 mM DTT.

Redox controlled dynamics of CrSBPase

Removed in the revised version: CrSBPase crystallized under reducing treatment (10 mM TCEP). Among the eight subunits in the asymmetric unit, C115-C120 disulfide is absent in two chains (subunits C colored in magenta and subunits E colored in salmon, figure 3D and 3F), while the other six subunits still have the C115-C120 disulfide bridge (figure 3B,C,E,G-I). All other cysteines are unaffected by the reducing treatment, consistently with their thiol form in crystal structure of oxidized state. These six subunits present mildly variable conformations of the A113SCAGTAC120 lasso conformation similar to the untreated oxidized structure of CrSBPase (figure 3A). Reduction of the C115-C120 disulfide increases local disorder, and the main chain electron density became uninterpretable over a few residues for subunits C (residues 115-118, 124-127), and subunits E (residues 119-127) (supplementary figure 4).

Removed in the revised version: Crystallographic structures of CrSBPase under reducing treatment: local disorder of the A113SCAGTAC120 loop. Main chain was traced according to crystallographic b-factor, with large orange sections representing high b-factor values and thin blue sections representing low b-factors. A-I. Aligned structures of CrSBPase protomers without redox treatment (A, 7b2o chain A) or in the presence of 10 mM TCEP reducing agent (B-I, 7zuv chains A-H).

Earlier version: Molecular dynamics (MD) simulations of the reduced structure were performed by starting from a dimer of the crystallized oxidized form and forcing C115 and C120 residues to be reduced in the thiol state (-SH). We then ran two independent replicas of 20 µs of MD. In addition, we also ran two independent replicas of 2 µs of MD in the oxidized form. In one of the replicas of the reduced form (MD1), the protein structure barely changed as observed in the SG(C115)/SG(C120) distance that plateaued at ∼7 Å, as well as in the RMSD with respect to the starting conformation that stayed around 2 Å (see supplementary figures 8-9 for plots and 12 for images of the structures). However, in the other replica of the reduced form (MD2), significant changes of conformation appeared: the SG(C115)/SG(C120) distance ended up at ∼10 Å with peaks at 20 Å and the RMSD values ended up at 3.7 and 2.8 Å (supplementary figures 8-9). The difference between MD1 and MD2 is stronger when one focuses on the RMSD of the mobile motif (residues 112 to 131, supplementary figure 10). We observed that after 17 µs of MD all chains adopted an equilibrated conformation and then fluctuated around that conformation (supplementary figure 10); representative structures extracted from MD were thus obtained from a clustering on the time frame (or block) 17-20 µs.

Revised version: CrSBPase crystallized under reducing treatment (10 mM TCEP). Among the eight subunits in the asymmetric unit, C115-C120 disulfide is absent in two chains (subunits C colored in magenta and subunits E colored in salmon, figure 3D and 3F), while the other six subunits still have the C115-C120 disulfide bridge (figure 3B,C,E,G-I). All other cysteines are unaffected by the reducing treatment, consistently with their thiol form in crystal structure of oxidized state. These six subunits present mildly variable conformations of the A113SCAGTAC120 lasso conformation similar to the untreated oxidized structure of CrSBPase (figure 3A). Reduction of the C115-C120 disulfide increases local disorder, and the main chain electron density became uninterpretable over a few residues for subunits C (residues 115-118, 124-127), and subunits E (residues 119-127) (supplementary figure 4). Molecular dynamics (MD) simulations of the reduced structure were performed by starting from a dimer of the crystallized oxidized form and forcing C115 and C120 residues to be reduced in the thiol state (-SH). We then ran two independent replicas of 20 µs of MD. In addition, we also ran two independent replicas of 2 µs of MD in the oxidized form. In one of the replicas of the reduced form (MD1), the protein structure barely changed as observed in the SG(C115)/SG(C120) distance that plateaued at ∼7 Å, as well as in the RMSD with respect to the starting conformation that stayed around 2 Å (see supplementary figures 8-9 for plots and 12 for images of the structures). However, in the other replica of the reduced form (MD2), significant changes of conformation appeared: the SG(C115)/SG(C120) distance ended up at ∼10 Å with peaks at 20 Å and the RMSD values ended up at 3.7 and 2.8 Å (supplementary figures 8-9). The difference between MD1 and MD2 is stronger when one focuses on the RMSD of the mobile motif (residues 112 to 131, supplementary figure 10). We observed that after 17 µs of MD all chains adopted an equilibrated conformation and then fluctuated around that conformation (supplementary figure 10); representative structures extracted from MD were thus obtained from a clustering on the time frame (or block) 17-20 µs.

Added in the revised version: Crystallographic structures of CrSBPase under reducing treatment: local disorder of the A113SCAGTAC120 loop.

Added in the revised version: Main chain was traced according to crystallographic b-factor, with large orange sections representing high b-factor values and thin blue sections representing low b-factors. A-I. Aligned structures of CrSBPase protomers without redox treatment (A, 7b2o chain A) or in the presence of 10 mM TCEP reducing agent (B-I, 7zuv chains A-H).

The two chains A and B from MD2 adopt a different final conformation (see figure 4A for the overlap with the crystallized oxidized structure and figure 4B for the overlap with the crystallized reduced structure), and we cannot directly conclude on which one is the most stable. Moreover, the residues 115 to 128 of each chain from MD1/2 do not overlap well with the crystallized reduced structure which means that neither of the two chains conformations from MD simulations has converged towards the crystallized state. However, we can observe in the reduced crystallized structure that residues 134-148 form an α-helix which is also present in chain B from MD2 whereas this structure is not kept in chain A from MD2 (see green chain on the right of figures 4A-B). In addition, mobile motif from chain B seems to open towards the solvent in a similar way with what was observed for FBPase (Chiadmi, Navaza et al. 1999). Thus, we conclude that chain B from MD2 is more representative of the true reduced conformation than chain A from MD2. If we compare the mobility of chain B from MD2 in the reduced state with one chain in the oxidized state (also chain B from MD2), we can see that the mobile motif is overall more flexible in the reduced state, which is especially true for residues around C120 (residues 118 to 123, supplementary figure 11). We then analyzed the RMSD of residues 112 to 131 from chain B during MD2 (supplementary figure 12): residues 112, 130 and 131 barely move which is expected since they belong or are directly connected to the two α-helixes flanking the mobile motif. At ∼6 µs of the trajectory almost all other residues (113-129 residues with the sole exception of C120) display an abrupt change of conformation, with a different magnitude though: starting from ∼2 Å, the RMSD jumped to 3 Å for some residues (residue 129) and to 15 Å for other residues (residues 125, 126) (supplementary figures 9, 10, and 12). At ∼13.5 µs, another change occurs which involves C120 but not all the other residues. Thus, we conclude that the opening of the mobile motif happens in a concerted way and not sequentially residue per residue.

Added in the revised version: Molecular dynamics simulation of CrSBPase after reduction.

Earlier version: Molecular dynamics simulation of CrSBPase after reduction. (A) Overlap of the crystallographic structure of oxidized CrSBPase and representative structures of equilibrated reduced CrSBPase during molecular dynamics simulation 2 (MD2). For structures extracted from MD, only residues 109 to 148 are displayed since most of the other residues are closely overlapping those of the crystallographic structure. (B) Overlap of the crystallographic structure of reduced CrSBPase and representative structures of equilibrated reduced CrSBPase during MD2. For structures extracted from MD, only residues 109 to 148 are displayed.

Revised version: (A) Overlap of the crystallographic structure of oxidized CrSBPase and representative structures of equilibrated reduced CrSBPase during molecular dynamics simulation 2 (MD2). For structures extracted from MD, only residues 109 to 148 are displayed since most of the other residues are closely overlapping those of the crystallographic structure. (B) Overlap of the crystallographic structure of reduced CrSBPase and representative structures of equilibrated reduced CrSBPase during MD2. For structures extracted from MD, only residues 109 to 148 are displayed.

Oligomeric states of CrSBPase

Previous studies reported a homodimeric state for SBPase from the moss Physcomitrella patens, deduced from the quaternary structure of recombinant protein in the asymmetric unit of the crystal (Gütle, Roret et al. 2016). Size-exclusion chromatography (SEC) profile of CrSBPase is polydisperse with at least three distinct oligomeric states (figure 6A). Notably, we observed that the protein elutes as a main peak of 79 kDa, close to the mass of a homodimer (theoretical mass from monomer sequence: 36,176.7 Da) and consistently with the crystallized state of moss enzyme. Thorough inspection of the elution profile of CrSBPase revealed a peak broadened by a later-eluting shoulder, with the apparent molecular mass of the tailing species at 48 kDa suggesting the presence of a monomeric state of CrSBPase in vitro. A third CrSBPase species elutes at an apparent molecular mass of 145 kDa, that fits with a homotetramer. According to the absorbance of each peak and considering that they represent all states of purified and oxidized CrSBPase in solution, we estimate a repartition percentage of species of 10:53:37 for the tetramer:dimer:monomer mixture (total = 100%). Re- injection of the dimer peak fraction over size-exclusion chromatography yielded a similar mixed profile, supporting a dynamic equilibrium over the time of the experiment (>2 h) (data not shown).

CrSBPase was crystallized in space groups P21212 or P1211 with respectively 6 and 8 protomers in the asymmetric unit. Proteins Interfaces Structures and Assemblies (PISA) analysis reveals that protomers of both crystals pack as homodimers with extensive subunit interface of buried area 4,170 Å² in a similar manner (figure 5A) (Krissinel and Henrick 2007). In the asymmetric unit of CrSBPase crystallized in reducing conditions (7zuv) or within crystallographic neighboring units of CrSBPase in oxidized state (7b2o), CrSBPase further packs into a homotetramer, a dimer of dimers (figure 5B). The apparent molecular mass of SBPase extracted from Chlamydomonas cultures cultivated in TAP medium under 40 µE.m-².s-1 illumination is comprised between 130 and 40 kDa with the homodimer representing the predominant state as supported by immunoblot analysis (figure 6B).

Added in the revised version: Oligomeric state of CrSBPase in the crystal.

Earlier version: Oligomeric state of CrSBPase in the crystal. (A) Asymmetric unit dimer of untreated CrSBPase (7b2o). Chains are represented in cartoon and colored cyan (chain B) and salmon (chain E). (B) Asymmetric unit homotetramer under reducing treatment (7zuv). Chains A (cyan), C (magenta), F (white), and H (orange) belong to the same asymmetric unit. (C) Close-up view on figure A homodimer interface. Loop 113-ASCAGTAC-120 from chain B (in cyan) is in 5-Å distance of neighboring chain E (in salmon). C115 and C120 are bonded by a disulfide bridge.

Revised version: (A) Asymmetric unit dimer of untreated CrSBPase (7b2o). Chains are represented in cartoon and colored cyan (chain B) and salmon (chain E). (B) Asymmetric unit homotetramer under reducing treatment (7zuv). Chains A (cyan), C (magenta), F (white), and H (orange) belong to the same asymmetric unit. (C) Close-up view on figure A homodimer interface. Loop 113-ASCAGTAC-120 from chain B (in cyan) is in 5-Å distance of neighboring chain E (in salmon). C115 and C120 are bonded by a disulfide bridge.

Added in the revised version: Oligomeric state of CrSBPase in vitro.

Earlier version: Oligomeric state of CrSBPase in vitro. (A) Size-exclusion chromatograms of CrSBPase wild-type (blue), mutant C115S (orange), mutant C120S (gray) on Superdex 200 26/600 GL column. (B) Size-exclusion fractionation of Chlamydomonas cell extracts. Chlamydomonas cell culture was harvested, lysed and the soluble fraction of the lysate was loaded on Superose6 16/600 size-exclusion column. Chromatography fractions were analyzed by western blot with anti-CrSBPase primary antibodies. First membrane was loaded with fractions eluted from 40 to 80 mL. Second membrane was loaded with fractions eluted from 80 to 120 mL. M lane is loaded with molecular mass standards ladder. Recombinant CrSBPase was loaded was loaded on the last lane to the right..

Revised version: (A) Size-exclusion chromatograms of CrSBPase wild-type (blue), mutant C115S (orange), mutant C120S (gray) on Superdex 200 26/600 GL column. (B) Size-exclusion fractionation of Chlamydomonas cell extracts. Chlamydomonas cell culture was harvested, lysed and the soluble fraction of the lysate was loaded on Superose6 16/600 size-exclusion column. Chromatography fractions were analyzed by western blot with anti-CrSBPase primary antibodies. First membrane was loaded with fractions eluted from 40 to 80 mL. Second membrane was loaded with fractions eluted from 80 to 120 mL. M lane is loaded with molecular mass standards ladder. Recombinant CrSBPase was loaded was loaded on the last lane to the right.

Earlier version: To further investigate the oligomeric dynamics of CrSBPase, we performed small-angle X-ray scattering (SAXS) analysis of purified (untreated) or reduced enzyme in vitro which reveals an increase in the protein radius of gyration from 39 to 62 Å upon reduction (supplementary figure 6). This observation implies a change in the oligomeric state of CrSBPase correlating with the redox state. As a matter of fact, purified CrSBPase tends to precipitate into amorphous aggregates that can be turned into droplets of liquid-liquid phase separation by treatment with 10 mM TCEP (supplementary figure 7). Since the LBH redox module is positioned at the homodimer interface (figure 5C), we postulated that the reduction of the C115-C120 disulfide bond and the subsequent local conformation changes depicted by MD affect the oligomer equilibrium of the enzyme. As a matter of fact, we observed that protein mutants C115S or C120S, which are unable to form the C115-C120 disulfide bond, elute in a different oligomer ratio than wild-type CrSBPase (figure 6A). Indeed, we observed a relative accumulation of monomer and tetramer species at the expense of the dimer state in both mutants. This suggests that C115-C120 is involved in the stabilization of the dimer. Meanwhile, the activity of the two Cys mutants still responded to DTT/Mg-dependent activation resembling the response of the wild-type protein. No activity was detected for native and untreated forms, while specific activities of C115S and C120S mutants upon treatment with DTT and MgSO4 (10 mM for both compounds) were respectively 96% and 107% with respect to wild-type maximum activity (13.3 ± 2.5 and 14,9 ± 0,9 µmol(NADPH) min-1 mg(SBPase)-1 for C115S and C120S mutants, respectively, supplementary figure 14). Based on these data, we can thus conclude that C115-C120 disulfide bond is not the sole thiol-dependent regulatory mechanism involved in the modulation of CrSBPase activity and that the redox dependency of CrSBPase catalysis likely depends on other cysteine pairs. Future studies are required to shed light on the identity of cysteine residues that contribute to CrSBPase catalytic modulation other than C115 and C120.

Revised version: To further investigate the oligomeric dynamics of CrSBPase, we performed small-angle X-ray scattering (SAXS) analysis of purified (untreated) or reduced enzyme in vitro which reveals an increase in the protein radius of gyration from 39 to 62 Å upon reduction (supplementary figure 6). This observation implies a change in the oligomeric state of CrSBPase correlating with the redox state. As a matter of fact, purified CrSBPase tends to precipitate into amorphous aggregates that can be turned into droplets of liquid-liquid phase separation by treatment with 10 mM TCEP (supplementary figure 7). Since the LBH redox module is positioned at the homodimer interface (figure 5C), we postulated that the reduction of the C115-C120 disulfide bond and the subsequent local conformation changes depicted by MD affect the oligomer equilibrium of the enzyme. As a matter of fact, we observed that protein mutants C115S or C120S, which are unable to form the C115-C120 disulfide bond, elute in a different oligomer ratio than wild-type CrSBPase (figure 6A). Indeed, we observed a relative accumulation of monomer and tetramer species at the expense of the dimer state in both mutants, as follows with tetramer:dimer:monomer ratio normalized on a scale of 100: WT=9:48:43; mutant C115S=42:25:33; mutant C120S=29:17:54. Size-exclusion chromatography suggests that C115-C120 is involved in the stabilization of the dimer. In order to test the hypothetical relevance of this disulfide bridge over catalysis, we assayed CrSBPase-mC115S and CrSBPase-m120S in the conditions described for wild-type enzyme in results section « CrSBPase enzymatic activity » above. Both Cys mutants behaved as wild-type enzyme, ie. they responded to DTT/Mg-dependent activation resembling that of the C115-C120 disulfide-containing protein. No activity was detected for native and untreated forms, while specific activities of C115S and C120S mutants upon treatment with DTT and MgSO4 (10 mM for both compounds) were near the specific activity of WT-SBPase (respectively 96% and 107% to the reference, 13.3 ± 2.5 and 14,9 ± 0,9 µmol(NADPH) min-1 mg(SBPase)-1 for C115S and C120S mutants) (supplementary figure 14). Considering the standard deviation of measured activities, there is no statistically significant difference between WT and mutants in our assays. Based on these data, we can thus conclude that C115-C120 disulfide bond is not the sole thiol-dependent regulatory mechanism involved in the modulation of CrSBPase activity and that the redox dependency of CrSBPase catalysis likely depends on one or several other cysteine pairs. Future studies are required to shed light on the identity of cysteine residues that contribute to CrSBPase catalytic modulation, other than C115 and C120.

Conclusions and perspectives

We report the first crystal structure of an SBPase from the microalga Chlamydomonas reinhardtii , a model for the molecular and cell biology of plants. SBPase is a photosynthetic enzyme involved in the regeneration phase of the CBBC cycle proved to be target of multiple regulations that we map onto our structural model. The native folding of CrSBPase is highly similar to that of SBPase from the moss Physcomitrium patens and to FBPase from chloroplast or cytoplasmic origins. Notwithstanding its propensity to form crystals, recombinant CrSBPase encompasses a range of oligomeric states, dominated by a homodimeric form susceptible to exchange for monomers, tetramers and higher- order assemblies. This oligomer equilibrium is dynamic and correlates with the presence of a disulfide bridge at the vicinity of the monomer-monomer interface, in the SBPase-specific LBH motif. We postulate that the reversible reduction/oxidation of the C115-C120 pair modulates the oligomerization of CrSBPase, through the induction of a local disorder that we observed in the crystal structure of the protein under reducing conditions and that correlates with computations of molecular dynamics.

How this oligomeric exchange impacts the actual activity of the enzyme is an open question but should be aligned with previous reports on the allosteric character of the structurally similar FBPase enzyme which follows dimer/tetramer exchanges (Barciszewski, Wisniewski et al. 2016) with a Hill cooperativity coefficient close to 2 representing positive cooperativity in the context of a homodimer (Giudici-Orticoni, Buc et al. 1990). Recent in vivo proteome mapping in Chlamydomonas reinhardtii localized SBPase in a punctate region surrounding the Rubisco pyrenoid nearby five other CBBC enzymes: phosphoglycerate kinase 1 (PGK), glyceraldehyde-3-phosphate dehydrogenases (GAP1, GAP3), fructose-bisphosphate aldolase (FBA3), ribulose-5-phosphate 3-epimerase (RPE1), and phosphoribulokinase (PRK) (Wang, Patena et al. 2022). Based on our structure-function study on recombinant SBPase, we can speculate that the altered redox state of SBPase may influence the movement to or from the liquid partition of the stroma, where SBPase is metabolically required to provide its contribution to CBBC reactions near other enzymes of the pathway.